rule_files:
- mise-prometheusRule.yaml
evaluation_interval: 1m
tests:
- interval: 1m
  input_series:
  - series: 'up{job="envoy-stats", namespace="mise"}'
    values: 1x0
  alert_rule_test:
  - eval_time: 1m
    alertname: MiseEnvoyScrapeDown
    exp_alerts:
    - exp_labels:
        severity: critical
      exp_annotations:
        summary: "Envoy scrape target down for namespace=mise"
        description: "Prometheus scrape for envoy-stats job in namespace mise is failing or missing."
