rule_files:
- arobit-prometheusRule.yaml
evaluation_interval: 1m
tests:
    - interval: 1m
  input_series:
  - series: 'up{job="arobit-forwarder",namespace="arobit"}'
    values: "0+0x15"
  alert_rule_test:
  - eval_time: 15m
    alertname: ArobitForwarderJobUp
    exp_alerts:
    - exp_labels:
        job: arobit-forwarder
        namespace: arobit
        alertname: ArobitForwarderJobUp
        severity: critical
      exp_annotations:
        description: |
          Arobit forwarder has not been reachable for the past 15 minutes.
          This may indicate that the Arobit forwarder is down, or experiencing a crash loop.
          Check the status of the Arobit forwarder pods, service endpoints, and network connectivity.
        runbook_url: TBD # https://runbooks.prometheus-operator.dev/runbooks/prometheus/prometheusdown
        summary: Arobit forwarder is unreachable for 15 minutes.
