rule_files:
- msi-credential-refresher-prometheusRule.yaml
evaluation_interval: 1m
tests:
- interval: 1m
  input_series:
  - series: 'credential_refresher_days_until_msi_credential_expiration_bucket{cluster="test-cluster",le="14"}'
    # Values need to increase over 30 minutes window for the increase() function to work
    # Starting from 0 and increasing by 1 every minute for 35 minutes
    values: '0+1x35'
  alert_rule_test:
  - eval_time: 35m
    alertname: ClusterCredentialExpiringSoon
    exp_alerts:
    - exp_labels:
        severity: critical
        cluster: test-cluster
        le: "14"
      exp_annotations:
        description: 'Cluster credential for cluster test-cluster is expiring in less than 30 days.'
        summary: 'Cluster credential expiring in less than 30 days'
        runbook_url: 'https://eng.ms/docs/cloud-ai-platform/azure-core/azure-cloud-native-and-management-platform/control-plane-bburns/azure-red-hat-openshift/azure-redhat-openshift-team-doc/doc/tsgs/credential-refresher-expiring-cert'
        correlationId: 'ClusterCredentialExpiringSoon/test-cluster'
